<problem display_name="Autocorrecteur" markdown="null" weight="10.0">
  <script type="loncapa/python">
def give_partial_credit(expect, ans):
    ans = str(ans)
    grade = 0
    if 'from thingz import *' in ans or 'import thingz' in ans or 'temperature()' in ans:
        grade += 1
    else:
        return {
            'input_list': [
                {'ok': False, 'msg': "Vous ne récupérez pas la valeur courante de la température",
    'grade_decimal': grade/10},
            ]
        }
    if 'open' in ans and 'w' in ans and 'data.txt' in ans:
        grade += 3
    else:
        return {
            'input_list': [
                {'ok': 'Partial', 'msg': "Vous n'ouvrez pas correctement le fichier data.txt",
    'grade_decimal': grade/10},
            ]
        }
    if 'deepsleep(' in ans:
        grade += 3
    else:
        return {
            'input_list': [
                {'ok': 'Partial', 'msg': "Vous avez oublié d'endormir le microcontrôleur en utilisant la fonction deepsleep()",
    'grade_decimal': grade/10},
            ]
        }
    if 'emoncms.fr/input/post' in ans or 'isconnected' in ans and 'Wokwi-GUEST' in ans:
        grade += 1
    else:
        return {
            'input_list': [
                {'ok': 'Partial', 'msg': "Vous n'auriez pas oublié d'envoyer la donnée vers emoncms.fr ? Indice : Le ssid doit être 'Wokwi-GUEST'",
    'grade_decimal': grade/10},
            ]
        }
    if 'write' in ans:
        grade += 2
    else:
        return {
            'input_list': [
                {'ok': 'Partial', 'msg': "Vous devez stocker la valeur courante de la température dans le fichier grâce à la fonction write()",
    'grade_decimal': grade/10},
            ]
        }
    if grade == 10:
        return {
            'input_list': [
                {'ok': True, 'msg': 'Votre note est de ' +
    str(grade) + '/10', 'grade_decimal': grade/10},
            ]
        }
</script>
  <customresponse cfn="give_partial_credit">
    <jsinput gradefn="getCode" height="640" width="100%" html_file="https://lms.fun-mooc.fr/asset-v1:MinesTelecom+04048+session02+type@asset+block@grader.html" title="Autocorrection MicroPython"/>
  </customresponse>
</problem>
