<problem display_name="Autocorrecteur" markdown="null" weight="2.0">
<p>
  Collez ci-dessous votre clé d'API <strong>« Lecture uniquement »</strong> et cliquez sur le bouton « Valider » ! L'aucorrecteur va se connecter à emoncms et vérifier que vous avez bien créer un compte et trois sources sur emoncms.fr. C'est à vous ⬇️
</p>    

<script type="loncapa/python">
import requests
from datetime import datetime, timedelta

def check_emoncms(expect,ans):
    if len(ans) != 32:
        return {
            'input_list': [
                {'ok': False, 'msg': 'Votre clé d\'API doit faire 32 caractères !'},
            ]
        }
    try:
        r = requests.get("https://emoncms.fr/input/list?apikey=" + str(ans))
        if r.status_code == 200:
            if len(r.json()):
                input_timestamp = int(r.json()[0]['time'])
                print(input_timestamp)
                if(datetime.now() - datetime.fromtimestamp(input_timestamp)) &lt; timedelta(minutes=10):
                    return {
                        'input_list': [
                            {'ok': True, 'msg': 'Bravo, au moins une source a été créée récemment !'},
                        ]
                    }
                else:
                    return {
                        'input_list': [
                            {'ok': False, 'msg': 'Aïe, vos sources ont été mises à jour il y a plus de 10 minutes. Veuillez réessayer après avoir mis à jour vos sources !'},
                        ]
                    }
            else:
                return {
                    'input_list': [
                        {'ok': False, 'msg': 'Il semblerait que vous n\'ayez pas créé de sources !'},
                    ]
                }
        else:
            return {
                'input_list': [
                    {'ok': False, 'msg': 'La requête HTTP a échouée !'},
                ]
            }
    except Exception as e:
        return {
            'input_list': [
                {'ok': True, 'msg': 'Vos sources ont été créées il y a plus de 10 minutes !'},
            ]
        }
</script>
  <customresponse cfn="check_emoncms">
    <textline size="42" label="Clé d'API lecture"/>
  </customresponse>
<solution>
  <div class="detailed-solution">
    <p>Bravo 👏 Votre clé d'API a permis de vérifier que vous avez bien créer un compte et au moins une source !</p>
  </div>
</solution>
</problem>
